Linguistics 531<br>
Fall 2023<br>
Jackson

## Things to remember about any homework assignment:

1. For this assignment, you will edit this jupyter notebook and turn it in. Do not turn in pdf files or separate `.py` files.
1. Late work is not accepted.
1. Given the way I grade, you should try to answer *every* question, even if you don't like your answer or have to guess.
1. You may *not* use `python` modules that we have not already used in class. (For grading, it needs to be able to run on my machine, and the way to do that is to limit yourself to the modules we've discussed and that are loaded into the Notebook.)
1. Don't use editors *other* than Jupyter Notebook to work on and submit your assignment, since they will mangle the autograding features: Google Colab, or even just editing the `.ipynb` file as a plain text file. Diagnosing and fixing that kind of problem takes a lot of my time, and that means less of my time to offer constructive feedback to you and to other students.
1. You may certainly talk to your classmates about the assignment, but everybody must turn in *their own* work. It is not acceptable to turn in work that is essentially the same as the work of classmates. Using someone else's code and simply changing variable or object names is *not* doing your own work.
1. All code must run. It doesn't have to be perfect, it may not do all that you want it to do, but it must run without error. Code that runs with errors will get no credit from the autograder.
1. Code must run in reasonable time. Assume that if it takes more than *5 minutes* to run (on your machine), that's too long.
1. Make sure to select `restart, run all cells` from the `kernel` menu when you're done and before you turn this in!

my name: Ashwin Raj

people I talked to about the assignment: *\<FILL IN HERE\>*

# Homework #1*

**This is due Tuesday, October 17, 2023 at noon (Arizona time).**

This assignment deals with the `NewB` corpus. We'll be using this is many of our later homeworks, as well. You should download it from [here](https://github.com/JerryWei03/NewB).

Before you start, here's a bit of background to get you headed in the right direction. In general contexts, the terms *file* and *document* can often be used interchangeably. However, in the following explanation, we follow the readings and the lecture in using them with these formal definitions:

*file* --> a unit on our computer in which data is stored

*document* --> an abstract unit of text in our corpus that we will be searching over, which we define based on the nature of our corpus and the kinds of results that we think will be most useful to users; one or more documents will be the units that we want to return to users as a result of a search query; the set of all documents is called the collection, and each document is described by a unique document ID.

One of the tasks that we need to do to prepare for searching a collection is to separate the documents in the collection from [metadata about those documents](https://bernardmarr.com/what-is-metadata-a-simple-explanation-of-what-everyone-should-know/#:~:text=Quite%20simply%3A%20metadata%20is%20data,web%20pages%2C%20content%20or%20spreadsheets.). In the video lectures, the documents that comprise our collection are blog posts, and the metadata for each post includes an author identifier, the author's age and gender, the author's interests, and the author's zodiacal sign. The posts for a single author are arranged in a single file. The metadata about those documents is found in the file names, and one of the first tasks that Mike does in the lectures is to write a function to read in the documents (ie, each post) and link each document with its metadata.

For this question on the homework, you need to think about these same issues for the NewB corpus. What are the documents that comprise this collection, that we will be searching over, and that we will be returning as the result of a query? What metadata do we have about each document? How is all of this stored in a file that we will read in, so that we can build an index? To answer this question, you may need to read the paper that describes the NewB corpus, which you can find in the Readings module. In this homework, you will see which of the files from the corpus we will be working with. Look at the docstrings for the two functions you need to build, and see what their arguments and their output will be like. Read the assert statements to see what tests will be performed on the output, and that will also give you clues about what data structure you should use.

**But first, make my autograding life easier (and your own notebook more likely to be graded correctly):**

In order for me to develop this assignment, and in order for me to **grade** your submission for this assignment, I need to be working with the right file that we load our corpus from. On my machine, that file has this path:

```python
# Path on my own machine, needed for GRADING
newbfile = '/home/ejackson1/Downloads/linguistics/NewB/train_orig.txt'

# ie, DON'T CHANGE THIS CELL, CHANGE THE ONE BELOW!
#  If you change *this* cell, the autograding is likely to break.
```

However, for **you** to work on your own code, you need to point this notebook to the path for this file on your own machine. *You should enter the path on your own machine in the editable code cell below,* then uncomment that line so the notebook works on your machine. This means that the second code cell will take precedence in assigning the value of the path to the corpus, and you can write your code to open that file without problems.

**HOWEVER, BEFORE YOU SUBMIT to D2L** comment out **your** path again. This means that when I run the code on my own machine, it'll have the path that ***I*** need, and it'll grade your notebook properly.

```python
# YOUR path
newbfile = '/home/ashwinr136/ling531/ling531/NewB/train_orig.txt'
```

**1.** Describe the general logic and contents of the corpus (1-2 paragraphs, 2 points)

<!-- grade_id: q1a, points: 2, grade: true, solution: true -->
The corpus contains sentences relating to Donald Trump. There is a label at the start which corresponds to where the information was sourced from. Following the label is the text data. The corpus contains a mix of news headlines, article snippets, and discussion points about Donald Trump.

**2.** Describe the file structure. Specifically tell me what you'll have to do to parse those files to produce reasonable data structures for searching. (1-3 paragraphs, 2 points)

<!-- grade_id: q2a, points: 2, grade: true, solution: true -->
The file is structured with an integer label, followed by whitespace, followed by the sentence to analyze, followed by a newline. Sometimes the sentence extends to the next line so it is important to make sure we don't go strictly line by line in the file as we may be omitting part of a sentence. We will have to remove the whitespace in between the ID and the text and then split the line so that we retain both the ID and the text.

**3.** Fill out the function below to parse the `train_orig.txt` file into lines. We'll be treating each sentence as a separate *document* in our *collection*.

There are different ways that you can read a file in Python. For a refresher, you may find [this tutorial](https://realpython.com/read-write-files-python/#reading-and-writing-opened-files) helpful, or you can always check the [official documentation](https://docs.python.org/3/tutorial/inputoutput.html#reading-and-writing-files). You can approach the problem here in different ways, but whatever method you use, just be sure you can pass the tests. (2 points total)

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
def getlines(filename):
    '''read in file and return the lines
    
    args:
        filename: name of the file
    returns:
        lines: a list of lines (with trailing newlines \n removed)
    '''
    # YOUR CODE HERE
    with open(filename, 'r') as file:
        lines = file.readlines()
        stripped = [line.rstrip('\n') for line in lines]
        return stripped
    
```

<!-- grade_id: q3t1, points: 1, grade: true, solution: false -->
```python
# test 3a, 1 pt
res = getlines(newbfile)
assert type(res) == list
```

<!-- grade_id: q3t2, points: 1, grade: true, solution: false -->
```python
# test 3b, 1 pt
res = getlines(newbfile)
assert len(res) == 253781
```

**4.** Now fill out the function below to **take the output of `getlines()` and parse it into a list of tuples**. Each tuple consists of the publication ID as an integer (in the interval 0 to 10, inclusive) and the sentence as a string. If you're not sure how to process what you read from the file, you can always look at the first part of what you've read in (for instance, using a slice), or you can open the input file in a text editor. (4 points total)

<!-- grade_id: q4a, points: 0, grade: false, solution: true -->
```python
def makeSentences(lines):
    '''convert the list of sentences into a list of tuples
    
    args:
        lines: the list of lines produced by getlines()
    
    returns:
        [(<publication_id>, "<sentence as a string>"), ...]
        a list of tuples of the form
            publication_id: integer
            sentence: string
    '''
    # YOUR CODE HERE

    result = []
    for line in lines:
        tab = line.find('\t')
        if tab != -1:
            pub_id = line[:tab]
            text = line[tab+1:]
            
            if pub_id.isdigit():
                pub_id = int(pub_id)
                result.append((pub_id, text))
    return result
```

<!-- grade_id: q4t1, points: 1, grade: true, solution: false -->
```python
# test 4a, 1pt
lines = getlines(newbfile)
sentences = makeSentences(lines)
assert type(sentences) == list
```

<!-- grade_id: q4t2, points: 1, grade: true, solution: false -->
```python
# test 4b, 1pt
assert len(sentences) == 253781
```

<!-- grade_id: q4t3, points: 1, grade: true, solution: false -->
```python
# test 4c, 1pt
from collections import Counter
doc_counts = Counter([id for id, sentence in sentences])
assert all([count == 23071 for id, count in doc_counts.items()])
```

<!-- grade_id: q4t4, points: 1, grade: true, solution: false -->
```python
# test 4d, 1pt
assert sentences[92294] == (4, 'if they criticize trump the president attacks them')
```

```python

```