Linguistics 531<br>
Fall 2023<br>
Jackson

## Things to remember about any homework assignment:

1. For this assignment, you will edit this jupyter notebook and turn it in. Do not turn in pdf files or separate `.py` files.
1. Late work is not accepted.
1. Given the way I grade, you should try to answer *every* question, even if you don't like your answer or have to guess.
1. You may *not* use `python` modules that we have not already used in class. (For grading, it needs to be able to run on my machine, and the way to do that is to limit yourself to the modules we've discussed and that are loaded into the Notebook.)
1. Don't use editors *other* than Jupyter Notebook to work on and submit your assignment, since they will mangle the autograding features: Google Colab, or even just editing the `.ipynb` file as a plain text file. Diagnosing and fixing that kind of problem takes a lot of my time, and that means less of my time to offer constructive feedback to you and to other students.
1. You may certainly talk to your classmates about the assignment, but everybody must turn in *their own* work. It is not acceptable to turn in work that is essentially the same as the work of classmates. Using someone else's code and simply changing variable or object names is *not* doing your own work.
1. All code must run. It doesn't have to be perfect, it may not do all that you want it to do, but it must run without error. Code that runs with errors will get no credit from the autograder.
1. Code must run in reasonable time. Assume that if it takes more than *5 minutes* to run (on your machine), that's too long.
1. Make sure to select `restart, run all cells` from the `kernel` menu when you're done and before you turn this in!

my name: Ashwin Raj

people I talked to about the assignment: *\<FILL IN HERE\>*

# Homework #7*

**This is due Tuesday, November 28, 2023 at noon (Arizona time).**

This assignment continues with the `NewB` corpus (downloadable [here](https://github.com/JerryWei03/NewB)).

imports:

```python
import re
import numpy as np
from math import isclose
from collections import Counter
```

**As before, this section is for autograding:**

What I need on my machine to properly grade this:

```python
# Path on my own machine, needed for GRADING
newbfile = '/home/ejackson1/Downloads/linguistics/NewB/train_orig.txt'

# ie, DON'T CHANGE THIS CELL, CHANGE THE ONE BELOW!
#  If you change *this* cell, the autograding is likely to break.
```

*In the editable cell below, enter the path on your own machine,* then uncomment that line so the notebook works on your machine.

**BEFORE YOU SUBMIT to D2L, remember to comment out *your* path again.**

```python
# YOUR path
#newbfile = '/home/ashwinr136/ling531/ling531/NewB/train_orig.txt'
```

**1.** Flesh out the function below that reads in the data and extract sentences in class #2 and class #8. (1 point)

You'll return a list of sentences in these two classes. Each sentence should be a single string. (You'll write a function to normalize them in question 2.)

You should be able to use concepts and code from last week and earlier in the course.

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python
def get28(filename):
    '''reads in newB data and extracts classes 2 and 8
    
    args:
        filename: location+name of newB file
    returns:
        s2: a list of sentences of class 2
        s8: a list of sentences of class 8
    '''
    # YOUR CODE HERE
    sentences = []

    with open(filename, 'r') as file:
        for line in file:
            parts = line.split('\t')
            if len(parts) == 2:
                publication_id, sentence = int(parts[0]), parts[1].strip()
                sentences.append((publication_id, sentence))
    s2 = [sentence for code, sentence in sentences if code == 2]
    s8 = [sentence for code, sentence in sentences if code == 8]
    return s2, s8
```

<!-- grade_id: q1t1, points: 1, grade: true, solution: false -->
```python
s2,s8 = get28(newbfile)

# test 1a, 1pt
assert len(s2) == len(s8) == 23071
```

**2.** Write a function to normalize and tokenize your sentences: convert anything besides upper or lower case letters, numbers, and the percent sign to space, and then split the string on white space. (2 points)

You should be able to reuse your code from Assignment 6 for this.

*This should now be very familiar. You don't need to overthink this. In a real application, we would probably want to do something fancy to handle Roman letters with accents and diacritics, like á and ñ, which **do** occur in the data sets we'll be using, but we're not doing anything fancy here. If it's not an upper case ASCII letter, or a lower case ASCII letter, or a digit, or the percent sign, convert it to space and then split on any white space. A few document IDs to check, to make sure you're properly normalizing, include 47602, 48327, 48677, and 56254. If you're working in Ubuntu, you can open the `train_orig.txt` file in the built-in text editor (gedit) and perform a RegEx-based "Find and Replace" search for a pattern like `[^a-zA-Z0-9%\s]`, in order to find examples of the kinds of characters your function will need to filter out, and the word tokenization that would result.*

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
def tokenize(s):
    '''
    converts anything that is not a letter (upper or lower
    case ASCII) or number to space, and tokenizes.
    
    args:
        s: a sentence as a string
    returns:
        a list of normalized word tokens as strings
    '''
    # YOUR CODE HERE
    normalized_text = re.sub(r'[^a-zA-Z0-9%]', ' ', s.lower())
    tokens = normalized_text.split()
    
    return tokens
```

<!-- grade_id: q2t1, points: 1, grade: true, solution: false -->
```python
s2sentences = [tokenize(s) for s in s2]
s8sentences = [tokenize(s) for s in s8]

# test 2a, 1pt
assert s8sentences[1214] == ['spurned', 'of', 'a', 'washington', 'dc', 'parade', 'president', 'trump',
 'said', 'he', 'now', 'plans', 'to', 'visit', 'paris', 'to', 'see', 'the', 'parade', 'down',
 'the', 'avenue', 'des', 'champs', 'lys', 'es', 'to', 'mark', 'the', 'centennial', 'of', 'the',
 'end', 'of', 'world', 'war', 'i', 'in', 'mid', 'november']
```

<!-- grade_id: q2t2, points: 1, grade: true, solution: false -->
```python
# test 2b, 1pt
assert len(s2sentences[2535]) == 52
```

**3.** Split the normalized and tokenized sentences in these two categories (using `s2sentences` and `s8sentences` that were output from question 2) into training and test sets where the *last* 1000 sentences comprise each test set and the rest make up the training sets. (3 points)

Pay attention to the variable names that are in the comments.

Again, this should seem very similiar to something you did last week.

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
#train2 = all but the last 1000 sentences
# YOUR CODE HERE
train2 = s2sentences[:-1000]

#test2 = the last 1000 sentences
# YOUR CODE HERE
test2 = s2sentences[-1000:]

#train8 = all but the last 1000 sentences
# YOUR CODE HERE
train8 = s8sentences[:-1000]

#test8 = the last 1000 sentences
# YOUR CODE HERE
test8 = s8sentences[-1000:]
```

<!-- grade_id: q3t1, points: 1, grade: true, solution: false -->
```python
# test 3a, 1pt
assert len(train2) == len(train8) == 22071 and len(test2) == len(test8) == 1000
```

<!-- grade_id: q3t2, points: 1, grade: true, solution: false -->
```python
# test 3b, 1pt
assert train8[15000][:5] == ['illustration', 'caption', 'sarah',
                             'ricegetty', 'and']
```

<!-- grade_id: q3t3, points: 1, grade: true, solution: false -->
```python
# test 3c, 1pt
assert test8[55][:5] == ['a', 'plan', 'being', 'considered', 'by']
```

**4.** Extract the set of all words in class 2, class 8, and in both classes together. (3 points)

You should include both the test and training sets for each class; that is, you can use `s2sentences` and `s8sentences` from question 2 as your input, not the separate test and training sets that were created as output in **Question 3**.

*Hint: [Set comprehensions](https://www.pythonforbeginners.com/basics/set-comprehension-in-python) may be helpful here. Also, once you have `vocab2` and `vocab8`, what's an easy way to create `vocab28`?*

<!-- grade_id: q4a, points: 0, grade: false, solution: true -->
```python
#vocab2 = set of words in class 2
# YOUR CODE HERE
vocab2 = set(word for sentence in s2sentences for word in sentence)

#vocab8 = set of all words in class 8
# YOUR CODE HERE
vocab8 = set(word for sentence in s8sentences for word in sentence)

#vocab28 = set of all words in class 2 and 8
# YOUR CODE HERE
vocab28 = vocab2 | vocab8
```

<!-- grade_id: q4t1, points: 1, grade: true, solution: false -->
```python
# This will obviously depend on your normalization function; if you're
#  having trouble getting these numbers, check that your normalization
#  function is really converting all the appropriate characters to space.

# test 4a, 1pt
assert len(vocab2) == 20701
```

<!-- grade_id: q4t2, points: 1, grade: true, solution: false -->
```python
# test 4b, 1pt
assert len(vocab8) == 21353
```

<!-- grade_id: q4t3, points: 1, grade: true, solution: false -->
```python
# test 4c, 1pt
assert len(vocab28) == 27944
```

**5.** Write a function that will create a multinomial Naive Bayes model from training data. (3 points)

The function will take two arguments: a list of normalized and tokenized sentences in some class (that is, `train2` or `train8` from question 3) plus a list of all possible words (in any class&mdash;`vocab28` from question 4). Look at the `assert` statements below to see how the function will behave.

The function will return a dictionary from each possible word to its log probability value in that training data. Use add-one smoothing to make sure all possible words have a non-zero (non-infinite in log space) value.

*Hint: Are you counting words? A `Counter()` may be a good option for this task. Note that in order to implement smoothing  with a `Counter()`, you may need to tackle this in two steps. Think about the two data structures that are inputs to this function and consider what each of them might help you with for this task. The Review & Mastery activity for this week may give you some hints.*

<!-- grade_id: q5a, points: 0, grade: false, solution: true -->
```python
def makeModel(sentences,vocabulary):
    '''
    returns a multinomial naive bayes model for a list
    of sentences and a set of possible words. Values are
    expressed as log probabilities.
    
    Uses add-one smoothing to make sure all possible words
    have a non-zero (non-infinite in log space) value.
    
    args:
        sentences: a list of training sentences
        vocabulary: a set of (all) possible words
    returns:
        a dictionary from possible words to that word's log probability
    '''
    # YOUR CODE HERE
    word_counts = Counter()
    total_words = 0
    
    for sentence in sentences:
        word_counts.update(sentence)
        total_words += len(sentence)
        
    vocab_size = len(vocabulary)
    class_probs = {}
    for word in vocabulary:
        word_count = word_counts[word]
        
        smoothed_prob = (word_count + 1) / (total_words + vocab_size)
        
        log_prob = np.log(smoothed_prob)
        class_probs[word] = log_prob

    return class_probs
```

<!-- grade_id: q5t1, points: 1, grade: true, solution: false -->
```python
model2 = makeModel(train2,vocab28)
model8 = makeModel(train8,vocab28)

# test 5a, 1pt
assert len(model2) == len(model8) == 27944
```

<!-- grade_id: q5t2, points: 1, grade: true, solution: false -->
```python
# test 5b, 1pt
# I get -11.313875043456596
assert isclose(model2['hats'],-11.3139,abs_tol=0.0001)
```

<!-- grade_id: q5t3, points: 1, grade: true, solution: false -->
```python
# test 5c, 1pt
# I get -11.455193945760294
assert isclose(model8['hats'], -11.4552, abs_tol=0.0001)
```

**6.** Fill out the function below to apply a pair of multinomial naive bayes models (that your function in question 5 created) to a list of normalized and tokenized sentences for testing (that is, we'll use `test2` and `test8` for this). (2 points)

The function reports the number of sentences that match the second model better.

You will *not* need the Bayes correction because there are an equal number of items in both classes.

*Hint: If you have a Naive Bayes model, how do you evaluate the predicted probability of a test document according to that model?*

<!-- grade_id: q6a, points: 0, grade: false, solution: true -->
```python
def applyModels(modelA,modelB,sentences):
    '''
    Determine which of two models fits the test data better.
    
    args:
        modelA, modelB: Naive Bayes models as created by makeModel()
        sentences: a list of normalized & tokenized sentences
            as created by tokenize()
    returns:
        the number of sentences that fit modelB better
    '''
    # YOUR CODE HERE
    better_model_count = 0

    for sentence in sentences:
        # Calculate the log probability of the sentence under both models
        probA = sum(modelA.get(word, 0.0) for word in sentence)
        probB = sum(modelB.get(word, 0.0) for word in sentence)

        # Compare probabilities and update the counter
        if probB > probA:
            better_model_count += 1

    return better_model_count
```

<!-- grade_id: q6t1, points: 1, grade: true, solution: false -->
```python
# Note that because we've written this function to retun the number
#  of sentences that fit *model B* better, we need to be sure that
#  the category for model B always matches the category of the test set.

res2 = applyModels(model8,model2,test2)
res8 = applyModels(model2,model8,test8)

# test 6a, 1pt
assert res2 == 729
```

<!-- grade_id: q6t2, points: 1, grade: true, solution: false -->
```python
# test 6b, 1pt
assert res8 == 645
```

**7.** Summarize the results of your multinomial Naive Bayes models and discuss. (2 points)

How well did the model do here? How did this model do, compared to the previous two models that we tried (namely, classifying solely by document length, and then classifying based on a normal statistical distribution)? Why do you think you're seeing this kind of performance, with this kind of model?

*Hint: If you're not sure how to approach this, re-watch the beginning of video 7.2 and the end of 7.4, paying attention to the assumptions that the multinomial Naive Bayes model makes about the documents within each class. What assumptions did our **previous** models make about the documents within each class? How well do those assumptions hold of this data?*

<!-- grade_id: q7a, points: 2, grade: true, solution: true -->
The multinomial bayes model based on train2 returns 729 sentences out of 1000 that fit test2 better than model 8, and the model based on train8 returns 645 sentences out of 1000 that fit test8 better than model2. This model had a higher accuracy on the test data compared to classifying solely by document length and by classifying based on a normal distribution. This could be because the Naive Bayes model takes into account the actual words in the document rather than following a probability distribution, which makes the model more complex and in this case results in higher accuracy.

```python

```